---
title: 'Chapter 1'
description: 'Chapter description goes here.'
free_preview: true
---

## Course 1

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.you can type words in here
但是中文來試看看可不可以

`@instructions`
This is an example exercise.you can type words in here
但是中文來試看看可不可以

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
