---
title: 'Chapter 1'
description: 'Chapter description goes here.'
free_preview: true
---

## Course 1

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.you can type words in here
但是中文來試看看可不可以

`@instructions`
This is an example exercise.you can type words in here
但是中文來試看看可不可以印出 Hello world

學習用
```
print("hello world")
```

`@hint`
印出“文字”

`@pre_exercise_code`
```{python}
# Pre-Exercise code
print()

```

`@sample_code`
```{python}
# Sample code

```

`@solution`
```{python}
# Solution
print("hello world")

```

`@sct`
```{python}
# SUBMISSION CORRECTNESS TESTS (SCT)
print("hello world")
```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: 5d693bdfd4
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
